# Lab 1

In today's lab, you'll learn how to:

1. navigate Jupyter notebooks (like this one);
2. write and evaluate some basic *expressions* in Python, the computer language of the course;


# 1. Jupyter Notebooks
This webpage is called a Jupyter Notebook. A notebook is a place to write programs and view their results, and also to write text.

##  Text cells
In a notebook, each rectangle containing text or code is called a *cell*.

Text cells (like this one) can be edited by double-clicking on them. They're written in a simple format called [Markdown](http://daringfireball.net/projects/markdown/syntax) to add formatting and section headings.  You don't need to learn Markdown, but you might want to.

After you edit a text cell, click the "run cell" button at the top that looks like ▶| or hold down `shift` + `return` to confirm any changes. (Try not to delete the instructions of the lab.)

```python

```

## Code cells
Other cells contain code in the Python 3 language. Running a code cell will execute all of the code it contains.

To run the code in a code cell, first click on that cell to activate it.  It'll be highlighted with a little green or blue rectangle.  Next, either press ▶| or hold down `shift` + `return`.

Try running this cell:

```python
print("Hello, World!") # output to console
```

Output:
```
Hello, World!
```

And this one:

```python
print("\N{WAVING HAND SIGN}, \N{EARTH GLOBE ASIA-AUSTRALIA}!") # output to console with emojis, \N is used to specify Unicode characters by name
```

Output:
```
👋, 🌏!
```

The fundamental building block of Python code is an expression. Cells can contain multiple lines with multiple expressions. When you run a cell, the lines of code are executed in the order in which they appear. Every `print` expression prints a line. Run the next cell and notice the order of the output.

```python
print("First this line is printed,")
print("and then this one.")
```

Output:
```
First this line is printed,
and then this one.
```

**Question** Change the cell above so that it prints out:

    First this line,
    then the whole 🌏,
    and then this one.

*Hint:* If you're stuck on the Earth symbol for more than a few minutes, try talking to a neighbor or a staff member. That's a good idea for any lab problem.

## Writing Jupyter notebooks
You can use Jupyter notebooks for your own projects or documents.  When you make your own notebook, you'll need to create your own cells for text and code.

To add a cell, click the + button in the menu bar.  It'll start out as a text cell.  You can change it to a code cell by clicking inside it so it's highlighted, clicking the drop-down box next to the restart (⟳) button in the menu bar, and choosing "Code".

**Question** Add a code cell below this one.  Write code in it that prints out:
   
    A whole new cell! ♪🌏♪

(That musical note symbol is like the Earth symbol.  Its long-form name is `\N{EIGHTH NOTE}`.)

Run your cell to verify that it works.

##  The Kernel
The kernel is a program that executes the code inside your notebook and outputs the results. In the top right of your window, you can see a circle that indicates the status of your kernel. If the circle is empty (⚪), the kernel is idle and ready to execute code. If the circle is filled in (⚫), the kernel is busy running some code.

Next to every code cell, you'll see some text that says `[ ]`. Before you run the cell, you'll see `[ ]`. When the cell is running, you'll see `[*]`. If you see an asterisk (\*) next to a cell that doesn't go away, it's likely that the code inside the cell is taking too long to run, and it might be a good time to interrupt the kernel (discussed below). When a cell is finished running, you'll see a number inside the brackets, like so: `[1]`. The number corresponds to the order in which you run the cells; so, the first cell you run will show a 1 when it's finished running, the second will show a 2, and so on.

You may run into problems where your kernel is stuck for an excessive amount of time, your notebook is very slow and unresponsive, or your kernel loses its connection. If this happens, try the following steps:
1. At the top of your screen, click **Kernel**, then **Interrupt**.
2. If that doesn't help, click **Kernel**, then **Restart**. If you do this, you will have to run your code cells from the start of your notebook up until where you paused your work.
3. If that doesn't help, restart your server. First, save your work by clicking **File** at the top left of your screen, then **Save and Checkpoint**. Next, click **Control Panel** at the top right. Choose **Stop My Server** to shut it down, then **Start My Server** to start it back up. Then, navigate back to the notebook you were working on. You'll still have to run your code cells again.

# Numbers

Quantitative information arises everywhere in data science. In addition to representing commands to print out lines, expressions can represent numbers and methods of combining numbers. The expression `3.2500` evaluates to the number 3.25. (Run the cell and see.)

```python
3.2500
```

Output:
```
3.25
```

Notice that we didn't have to `print`. When you run a notebook cell, if the last line has a value, then Jupyter helpfully prints out that value for you. However, it won't print out prior lines automatically.

```python
print(2)
3
4
```

Output:
```
2
4
```

Above, you should see that 4 is the value of the last expression, 2 is printed, but 3 is lost forever because it was neither printed nor last.

You don't want to print everything all the time anyway.  But if you feel sorry for 3, change the cell above to print it.

## Arithmetic
The line in the next cell subtracts.  Its value is what you'd expect.  Run it.

```python
3.25 - 1.5
```

Output:
```
1.75
```

Many basic arithmetic operations are built into Python. The common operator that differs from typical math notation is `**`, which raises one number to the power of the other. So, `2**3` stands for $2^3$ and evaluates to 8.

The order of operations is the same as what you learned in elementary school, and Python also has parentheses.  For example, compare the outputs of the cells below. The second cell uses parentheses for a happy new year!

```python
6+6*5-6*3**2*2**3/4*7 # order of operations, PEMDAS
```

Output:
```
-720.0
```

```python
6+(6*5-(6*3))**2*((2**3)/4*7)
```

Output:
```
2022.0
```

In standard math notation, the first expression is

$$6 + 6 \times 5 - 6 \times 3^2 \times \frac{2^3}{4} \times 7,$$

while the second expression is

$$6 + (6 \times 5 - (6 \times 3))^2 \times \left(\frac{\left(2^3\right)}{4} \times 7\right).$$

**Question ** Write a Python expression in this next cell that's equal to $\displaystyle 5 \times \left(3 \frac{10}{11}\right) - 50 \frac{1}{3} + 2^{.5 \times 22} - \frac{7}{33} + 5$.  That's five times three and ten elevenths, minus fifty and a third, plus two to the power of half twenty-two, minus seven thirty-thirds plus five.  By "$\displaystyle 3 \frac{10}{11}$" we mean $\displaystyle 3+\frac{10}{11}$, not $\displaystyle 3 \times \frac{10}{11}$.

Replace the ellipses (`...`) with your expression.  Try to use parentheses only when necessary.

*Hint:* The correct output should start with a familiar number.

```python
...
```

Output:
```
Ellipsis
```

```python
ten = 3 * 2 + 4 # assigning value to variable
```

When you run that cell, Python first computes the value of the expression on the right-hand side, `3 * 2 + 4`, which is the number 10.  Then it assigns that value to the name `ten`.  At that point, the code in the cell is done running.

After you run that cell, the value 10 is bound to the name `ten`:

```python
ten
```

Output:
```
10
```

The statement `ten = 3 * 2 + 4` is not asserting that `ten` is already equal to `3 * 2 + 4`, as we might expect by analogy with math notation.  Rather, that line of code changes what `ten` means; it now refers to the value 10, whereas before it meant nothing at all.

If the designers of Python had been ruthlessly pedantic, they might have made us write

    define the name ten to hereafter have the value of 3 * 2 + 4

instead.  You will probably appreciate the brevity of "`=`"!  But keep in mind that this is the real meaning.

**Question ** Run the following cell which uses a variable name `eleven` that hasn't been assigned to anything. You'll see an error!

A common pattern in Jupyter notebooks is to assign a value to a name and then immediately evaluate the name in the last line in the cell so that the value is displayed as output.

```python
close_to_pi = 355/113
close_to_pi
```

Output:
```
3.1415929203539825
```

Another common pattern is that a series of lines in a single cell will build up a complex computation in stages, naming the intermediate results.

```python
semimonthly_salary = 842.5
monthly_salary = 2 * semimonthly_salary # assigning value to variable using another variable
number_of_months_in_a_year = 12
yearly_salary = number_of_months_in_a_year * monthly_salary # operation using variables
yearly_salary # output the value of yearly_salary to console
```

Output:
```
20220.0
```

Names in Python can have letters (upper- and lower-case letters are both okay and count as different letters), underscores, and numbers.  The first character can't be a number (otherwise a name might look like a number).  And names can't contain spaces, since spaces are used to separate pieces of code from each other.

Other than those rules, what you name something doesn't matter *to Python*.  For example, this cell does the same thing as the above cell, except everything has a different name:

```python
a = 842.5
b = 2 * a
c = 12
d = c * b
d
```

Output:
```
20220.0
```

```python
abs(5)
```

Output:
```
5
```

```python
abs(21.7 - 18.2) # positive difference between two numbers
```

Output:
```
3.5
```

Tests for q3_1_2 (points: None)
```
>>> # It looks like you didn't give anything the name
>>> # seconds_in_a_decade.  Maybe there's a typo, or maybe you
>>> # just need to run the cell below Question 3.2 where you defined
>>> # seconds_in_a_decade.  (Click that cell and then click the "run
>>> # cell" button in the menu bar above.)
>>> 'seconds_in_a_decade' in vars()
True
```
```
>>> # It looks like you didn't change the cell to define
>>> # seconds_in_a_decade appropriately.  It should be a number,
>>> # computed using Python's arithmetic.  For example, this is
>>> # almost right:
>>> #   seconds_in_a_decade = 10*365*24*60*60
>>> seconds_in_a_decade != ...
True
```
```
>>> # It looks like you didn't account for leap years.
>>> # There were 2 leap years and 8 non-leap years in this period.
>>> # Leap years have 366 days instead of 365.
>>> seconds_in_a_decade != 315360000
True
```
```
>>> seconds_in_a_decade == 315532800
True
```


Tests for q3_3_1 (points: None)
```
>>> # Fill in the row
>>> #   time = ...
>>> # with something like:
>>> #   time = 4.567
>>> # (except with the right number).
>>> time != ...
True
```
```
>>> # Read the text above the question to see what
>>> # time should be. 
>>> round(time, 5)
1.2
```
```
>>> # Fill in the row
>>> #   estimated_distance_m = ...
>>> # with something like:
>>> #   estimated_distance_m = 4.567
>>> # (except with the right number). 
>>> estimated_distance_m != ...
True
```
```
>>> # Note that the units are meters, but the text used
>>> # centimeters.
>>> estimated_distance_m != 113
True
```
```
>>> # Read the text above the question to see what
>>> # estimated_distance_m should be.
>>> round(estimated_distance_m, 5)
1.13
```


Tests for q3_3_2 (points: None)
```
>>> # Fill in the line that currently says
>>> #   predicted_distance_m = ...
>>> # in the cell above.
>>> predicted_distance_m != ...
True
```
```
>>> # Compute predicted_distance_m using the formula in the text
>>> # above.  Hint: it should start with something like this:
>>> #   predicted_distance_m = (1/2) * gravity_constant ...
>>> round(predicted_distance_m, 5)
1.17022
```
```
>>> round(difference, 5)
0.04022
```


Tests for q4_1_1 (points: None)
```
>>> # Fill in the line
>>> #   num_avenues_away = ...
>>> # in the cell above. 
>>> num_avenues_away != ...
True
```
```
>>> # Remember to compute the absolute value of 7-10.  Traveling 
>>> # "-3 blocks" doesn't really make sense!
>>> num_avenues_away != -3
True
```
```
>>> num_avenues_away
3
```
```
>>> manhattan_distance
1462
```


Tests for q51 (points: None)
```
>>> round(genghis_distance_from_average_in, 3)
1.5
```


Tests for q5_1_1 (points: None)
```
>>> round(min_length_difference, 5)
3.9
```
